```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment4.ipynb")
```

# Clustering in Practice and K-Means Clustering

Please complete the following assignment. Run all cells and submit your completed notebook through Gradescope.

- Clustering is a fundamental unsupervised learning technique used to group similar data points based on their features. It helps in discovering patterns and structures within datasets without predefined labels.

- The k-means algorithm aims to partition a dataset into k clusters, assigning each data point to the cluster whose centroid is closest to that point.

**In this assignment, the aim is to check and reinforce your understanding  of clustering, specifically K-Means clustering.**


**Note:** Do not change any of the parameters given in the starter code (e.g., random_state=42, number of samples, centers, or cluster standard deviation).  

These fixed values ensure your results are reproducible and match the expected test cases.  

<!-- BEGIN QUESTION -->

#### 1. Which of the following best describes the objective function minimized in k-means clustering?

A) The variance between cluster centroids.

B) The sum of squared distances between all pairs of data points.

C) The sum of squared distances between each point and its assigned cluster centroid.

D) The determinant of the covariance matrix of the clusters.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### 2. Which of the following is a known limitation of k-means?

A) It cannot be implemented in higher dimensions.

B) It performs poorly when clusters have non-spherical shapes.

C) It can only work with datasets having fewer than 1,000 points.

D) It guarantees the global optimum clustering solution.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### 3. What is the main advantage of using the k-means++ initialization method over the standard k-means initialization?

A) It selects initial centroids that are maximally distant from each other, ensuring convergence to the global optimum.

B) It reduces the computational complexity of k-means from quadratic to linear time with respect to the number of data points.

C) It allows k-means clustering to automatically determine the optimal number of clusters without prior specification.

D) It probabilistically selects initial centroids based on data density, leading to faster convergence and improved clustering results.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### 4-a. Clustering on Isotropic Gaussian Distributed Data

- A synthetic dataset has been generated using sklearn’s `make_blobs`.  
- Apply **K-Means clustering** with `k = 3`.  
- Plot the clustered data and highlight the cluster centers with a distinct marker.  
- Return:  
  - The cluster centers  
  - The predicted cluster labels  
  - The Within-Cluster Sum of Squares (WCSS)  

```python
import numpy as np
import matplotlib.pyplot as plt
from sklearn.datasets import make_blobs
from sklearn.cluster import KMeans

def ans4a():
    X, y_true = make_blobs(n_samples=300, centers=3, cluster_std=1.0, random_state=42)
    ...

    # Return the following values:
    # - The cluster centers
    # - The predicted cluster labels
    # - The Within-Cluster Sum of Squares (WCSS)
    # - The original data points
    return centers, y_kmeans, inertia, X


centers, y_kmeans, inertia, X = ans4a()
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

#### 4-b. Plot the original data points and the clustered data points
 
- Plot the clustered data and highlight the cluster centers with a distinct marker.  

```python
...
```

<!-- END QUESTION -->

#### 5a. Clustering on funky-looking anisotropic data

- A synthetic dataset has been generated using sklearn’s `make_moons`.  
- Apply **K-Means clustering** with `k = 2`.  
- Return:  
  - The predicted cluster labels  
  - The true labels  
  - The Adjusted Rand Index (ARI)
  - The original data points
  - The true labels
  - The cluster centers

```python
from sklearn.datasets import make_moons
from sklearn.cluster import KMeans
from sklearn.metrics import adjusted_rand_score

def ans5a():
    # Generate noisy moons dataset
    X, y_true = make_moons(n_samples=300, noise=0.15, random_state=42)
    ...

    # Return the following values:
    # - The predicted cluster labels
    # - The true labels
    # - The Adjusted Rand Index (ARI)
    # - The original data points
    # - The true labels
    # - The cluster centers
    return y_kmeans, y_true, ari, X, y_true, centers_


y_kmeans, y_true, ari, X, y_true, centers = ans5a()
```

```python
grader.check("q5a")
```

<!-- BEGIN QUESTION -->

#### 5-b. Plot the clustered data

- Make two plots:  
  - Plot the original data points.
  - Plot the clustered data and highlight the cluster centers with a distinct marker.  


```python
# Plot the original, ground truth data points.

...
```

```python
# Now plot the k-means clustering results with the cluster centers highlighted.

...
```

<!-- END QUESTION -->

#### 6-a. Comparing K-Means Initialization Methods

- In this question, you will **compare** K-Means clustering using **random initialization** and **k-means++ initialization** on synthetic data generated with `make_blobs`.  
- Apply K-Means with the same number of clusters for both initialization methods.
- Compare the two approaches by reporting:  
  - The number of iterations until convergence  
  - The WCSS (within-cluster sum of squares)    
- Return the following values:
  - `n_iter_random`, `n_iter_kpp`, `wcss_random`, `wcss_kpp`

```python
import warnings

def ans6a():
    # Generated synthetic dataset with 5 clusters
    X, y_true = make_blobs(n_samples=1000, centers=5, cluster_std=0.60, random_state=42)
    ...

ans6a()
```

```python
grader.check("q6a")
```

<!-- BEGIN QUESTION -->

#### 6-b. Briefly explain how the choice of initialization (random vs. k-means++) influences the convergence speed and the quality of the clustering results.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### 7-a. Determining Optimal K Using the Elbow Method  

- In this question, you will explore how to select the optimal number of clusters using the **Elbow Method**.  
- Run **K-Means clustering** on a synthetic dataset for different values of `k` (from 1 to 10), using the **k-means++ initialization** strategy.  
- For each `k`:  
  - Record the Within-Cluster Sum of Squares (WCSS).  
  - Record the Rand Index score relative to the true labels.  
- Based on the results, determine the **optimal `k` value**.  
- Return the following values (in this order):  
  - `optimal_k`: the optimal number of clusters, 
  - `rand_scores_list`: the list of Rand Index scores, 
  - `wcss_list`: the list of WCSS values, 
  - `k_values`: the list of k values, 
  - `wcss_list`: the list of WCSS values 

```python
from sklearn.datasets import make_blobs
from sklearn.cluster import KMeans
from sklearn.metrics import adjusted_rand_score
import warnings

def ans7a():
    # Generated synthetic dataset with 5 clusters
    X, y_true = make_blobs(n_samples=1000, centers=5, cluster_std=0.60, random_state=42)
    ...

optimal_k, rand_scores, wcss, k_values, wcss_list = ans7a()
```

```python
grader.check("q7a")
```

<!-- BEGIN QUESTION -->

#### 7-b. Elbow Curve

- Create a visualization of the **Elbow Curve** (WCSS vs. k). 

```python
...
```

<!-- END QUESTION -->



Tests for q4a (points: 2)
```
>>> def test_ans4a(ans4a):
...     expected_centers = np.array([[-2.63, 9.04], [-6.88, -6.98], [4.75, 2.01]])
...     expected_wcss = 566.86
...     (centers, labels, wcss, X) = ans4a()
...     assert centers.shape == (3, 2)
...     assert labels.shape == (300,)
...     assert set(labels) == {0, 1, 2}
...     assert np.allclose(centers, expected_centers, atol=0.1)
...     assert np.isclose(wcss, expected_wcss, atol=1)
>>> test_ans4a(ans4a)
```


Tests for q5a (points: 2)
```
>>> def test_ans5a(ans5a):
...     (pred, true, ari, X, y_true, centers) = ans5a()
...     assert len(pred) == 300
...     assert len(true) == 300
...     assert set(pred) == {0, 1}
...     assert set(true) == {0, 1}
...     assert 0 <= ari <= 1
...     assert np.isclose(ari, 0.28, atol=0.05)
>>> test_ans5a(ans5a)
```


Tests for q6a (points: 3)
```
>>> import warnings
>>> def test_ans6a(ans6a):
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore')
...         (n_iter_random, n_iter_kpp, wcss_random, wcss_kpp) = ans6a()
...     assert n_iter_random == 6
...     assert n_iter_kpp == 2
...     assert abs(wcss_random - 698.77) < 1
...     assert abs(wcss_kpp - 698.77) < 1
```


Tests for q7a (points: 2)
```
>>> import warnings
>>> def test_ans7a(ans7a):
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore')
...         (k, rand_scores, wcss, k_values, wcss_list) = ans7a()
...     assert isinstance(k, int)
...     assert len(rand_scores) == 10 and len(wcss) == 10
...     assert k == 5
...     assert np.allclose(wcss, [57466, 27529, 6145, 1891, 698, 644, 591, 537, 503, 451], atol=10)
...     assert wcss[0] > wcss[-1]
...     assert wcss[0] > 57000
...     assert np.allclose(rand_scores[1:], [0.21, 0.61, 0.78, 0.99, 0.93, 0.86, 0.78, 0.7, 0.61], atol=0.05)
...     assert wcss[0] / wcss[1] > 2
...     assert wcss[2] / wcss[3] > 3
...     assert wcss[8] - wcss[9] > 50
```
