```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment3.ipynb")
```

# Linear Algebra Refresher and an Introduction to NumPy

This assignment aims to refresh your understanding of key linear algebra concepts that are fundamental to data science and machine learning, with a focus on practical implementation using NumPy. NumPy (Numerical Python) is a foundational library for scientific computing in Python and is another vital tool we'll be using in this class. It provides support for large, multi-dimensional arrays and matrices, along with a vast collection of high-level mathematical functions to operate on these arrays efficiently. 

You'll find the [Numpy documentation](https://numpy.org/doc/stable/) useful in this assignment.

You should write code in the cells marked with "YOUR CODE HERE" or "..."

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd

np.random.seed(701)
```

## Part 1 - Basics

You can turn a Python array, such as a list, of any dimensions into a Numpy array using `np.array()`.

```python
py_array = [1, 2, 3, 4, 5]
np_array = np.array(py_array)

py_array2 = [[1, 2, 3], [4, 5, 6]]
np_array2 = np.array(py_array2)

print(py_array)
print(np_array)
print(py_array2)
print(np_array2)
```

#### **Question 1:** 

A very common operation is checking the shape of an array. You can check the shape of a numpy array using the `.shape` attribute. Print the shapes of the two numpy arrays we created above in the next cell.

```python
np_array_shape = ...
print(np_array_shape)
print(type(np_array_shape))

np_array2_shape = ...
print(np_array2_shape)
```

```python
grader.check("q1")
```

#### **Question 2:**

Now, use the [numpy random module](https://numpy.org/doc/stable/reference/random/index.html) to create a 3x3 array of random integers from 1 to 10. 

```python
rng = np.random.default_rng()
random_array = ...

print("3x3 array of random numbers from 1 to 10:")
print(random_array)
```

```python
grader.check("q2")
```

You can perform element wise addition, subtraction, multiplication, and division on two arrays. 

```python
a1 = np.array([1, 2, 3])
a2 = np.array([4, 5, 6])

print("Element wise addition:", a1 + a2)
print("Element wise subtraction:", a1 - a2)
print("Element wise multiplication:", a1 * a2)
print("Element wise division:", a1 / a2)
```

#### **Question 3:** 

Use the [`.dot` method](https://numpy.org/doc/1.21/reference/generated/numpy.ndarray.dot.html) to compute the dot product of vectors $a_1$ and $a_2$.

```python
dot_product = ...
```

```python
grader.check("q3")
```

### Question 4

Compute the matrix product of the two matrices defined below, $M_1$ and $M_2$. You can use [`np.matmul`](https://numpy.org/doc/stable/reference/generated/numpy.matmul.html) or its shorthand `@`. 

```python
M1 = np.array([[1, 2], 
               [3, 4]])

M2 = np.array([[5, 6], 
               [7, 8]])

product = ...

print("Matrix product of M1 and M2:")
print(product)
```

```python
grader.check("q4")
```

## Broadcasting

A key strength of Numpy is [broadcasting](https://numpy.org/doc/stable/user/basics.broadcasting.html) - treating arrays of different shapes during arithmetic operations. Let's play with it here. 

```python
# EXAMPLE
a = np.array([1.0, 2.0, 3.0])
b = 2.0

print("Broadcasting a + b:", a + b)
print("Broadcasting a - b:", a - b)
print("Broadcasting a * b:", a * b)
print("Broadcasting a / b:", a / b)

```

### Question 5

For this exercise, do the following:

1. Create a 2D array called `a1` with shape (3, 4) containing random integers between 1 and 10 (exclusive), in other words between 1 and 9 (inclusive).
   
2. Create a 1D array called `b1` with shape (4,) containing random integers between 1 and 5 (exclusive), in other words between 1 and 9 (inclusive).
   
3. Use broadcasting to add b to each row of `a1`.
   
4. Print the original arrays and the result. 

```python
rng = np.random.default_rng(701)

a1 = ...
b1 = ...

print("A1:\n", a1)
print("B1:\n", b1)

result = ...

print("\nResult:\n", result)
```

```python
grader.check("q5")
```

### Question 6

Let's demonstrate broadcasting with a concrete example. We'll create a 100x100x3 array that will represent an RGB image (red, green, blue). Specifically, every row and column contains an array of 3 values - level of red (R), level of green (G) and level of blue (B). Every one of those values ('levels') is in the range 0 (black) to 255 (brightest).

```python
np.random.seed(701)

random_image = np.random.randint(0, 255, size=(100, 100, 3))

print("Here is a random 100x100x3 image:")
plt.imshow(random_image)
plt.axis('off')
plt.show()
```

```python
print("Level of red (R), level of green (G) and level of blue (B) at the top left corner of the image:")
print(random_image[0][0])
```

We can use broadcasting to alter the colors of the image. For this exercise, create a 1D array of length 3 to represent brightness adjustments. For example, [1.1, 0.9, 1.2] to brighten red, darken green, and brighten blue (NOTE: Do NOT use numpy.random for this, create it yourself.)

Then, use broadcasting to apply these adjustments to the entire image. 

You might need to clip the resulting values to ensure they stay within the valid range (0-255).

```python
from copy import deepcopy

changed_image = deepcopy(random_image)

brightness_adjustments = ...

changed_image = ...

changed_image = changed_image.astype(int)
plt.imshow(changed_image)
plt.axis('off')
plt.show()
```

Feel free to play around with the brightness adjustments! 

```python
grader.check("q6")
```

## Part 2 - Linear Algebra with Numpy

For this part, you should find the appropriate functions and methods from the documentation or online. 

### Question 7

Calculate the $\ell_2$ norm of the vector $v_1$.  See numpy's [norm](https://numpy.org/doc/2.3/reference/generated/numpy.linalg.norm.html).

```python
np.random.seed(701)

v1 = np.random.randint(0, 10, size=10)
print("Vector v1:", v1)

magnitude = ...

print(f"l2 norm of v1: {magnitude:.2f}")
```

```python
grader.check("q7")
```

### Matrix Decompositions

### Question 8

Use Numpy to get the eigenvalues and eigenvectors of a given matrix $M_3$. 
See numpy's [eig](https://numpy.org/doc/2.3/reference/generated/numpy.linalg.eig.html) function.

```python
M3 = np.array([[10, 15, 20],
               [25, 30, 35],
               [40, 45, 50]])

eigenvalues, eigenvectors = ...

print(f"Eigenvalues of M3: {[f'{x:.2f}' for x in eigenvalues]}")
print(f"Eigenvectors of M3: {[[f'{x:.2f}' for x in v] for v in eigenvectors]}")
```

```python
grader.check("q8")
```

#### Question 9

Verify that $Av = λv$ for each eigenpair $(λ, v)$ you found. 

Reconstruct the matrix using the eigendecomposition that we describe in the
[notes](https://tools4ds.github.io/DS701-Course-Notes/04-Linear-Algebra-Refresher.html#eigendecomposition). 


```python
M3_reconstructed = ...

print("Reconstructed M3:", M3_reconstructed)
```

```python
grader.check("q9")
```

### Question 10: Singular Value Decomposition

Decompose and reconstruct the matrix in the same manner for the Singular Value Decomposition below. 
See [svd](https://numpy.org/doc/2.3/reference/generated/numpy.linalg.svd.html).

```python
U, s, V = ...

print(f"Singular values of M3: {[f'{x:.2f}' for x in s]}")
print(f"U matrix of M3: {[[f'{x:.2f}' for x in v] for v in U]}")
print(f"V matrix of M3: {[[f'{x:.2f}' for x in v] for v in V]}")

M3_reconstructed = ...
print("Reconstructed M3:", M3_reconstructed)
```

```python
grader.check("q10")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> def test_q1(np_array_shape, np_array2_shape):
...     assert np.all(np_array_shape == (5,)), f'Expected shape (5,), but got {np_array_shape}'
...     assert np.all(np_array2_shape == (2, 3)), f'Expected shape (2, 3), but got {np_array2_shape}'
>>> test_q1(np_array_shape, np_array2_shape)
```


Tests for q10 (points: 4)
```
>>> def test_svd(M3):
...     assert np.allclose(M3, U @ np.diag(s) @ V, atol=0.05), 'SVD reconstruction failed'
>>> test_svd(M3)
```


Tests for q2 (points: 1)
```
>>> def test_q2(random_array):
...     assert random_array.shape == (3, 3), f'Expected shape (3, 3), but got {random_array.shape}'
...     assert np.all(random_array >= 1) and np.all(random_array <= 10), f'Expected all elements to be in the range 1 to 10, but found elements outside this range.'
>>> test_q2(random_array)
```


Tests for q3 (points: 2)
```
>>> def test_q3(a1, a2):
...     assert np.allclose(dot_product, 32), f'Expected dot product {32}, but got {dot_product}'
>>> test_q3(a1, a2)
```


Tests for q4 (points: 2)
```
>>> def test_q4(M1, M2):
...     assert np.allclose(product, [[19, 22], [43, 50]]), f'Expected matrix product [[19, 22], [43, 50]], but got {product}'
>>> test_q4(M1, M2)
```


Tests for q5 (points: 3)
```
>>> def test_q5(a1, b1, result):
...     assert np.allclose(result, [[2, 10, 8, 11], [4, 9, 2, 11], [6, 10, 4, 5]]), f'Expected result [[ 2, 10,  8, 11], [ 4,  9,  2, 11], [ 6, 10,  4,  5]], but got {result}'
>>> test_q5(a1, b1, result)
```


Tests for q6 (points: 2)
```
>>> def test_q6(random_image, brightness_adjustments, changed_image):
...     assert np.allclose(changed_image, random_image * brightness_adjustments, atol=1), f"Didn't get expected changed image."
>>> test_q6(random_image, brightness_adjustments, changed_image)
```


Tests for q7 (points: 2)
```
>>> def test_q7(v1, magnitude):
...     assert np.isclose(magnitude, 18.76, atol=0.05), f'Expected magnitude {18.76}, but got {magnitude:.2f}'
>>> test_q7(v1, magnitude)
```


Tests for q8 (points: 2)
```
>>> def test_eigenvalues(M3, eigenvalues):
...     assert np.allclose(eigenvalues, [94.75, -4.75, 0.0], atol=0.05), f'Incorrect eigenvalues.'
>>> test_eigenvalues(M3, eigenvalues)
```
```
>>> def test_eigenvectors(M3, eigenvectors):
...     assert np.allclose(eigenvectors, [[-0.28, -0.78, 0.41], [-0.54, -0.07, -0.82], [-0.79, 0.63, 0.41]], atol=0.03), f'Incorrect eigenvectors'
>>> test_eigenvectors(M3, eigenvectors)
```


Tests for q9 (points: 2)
```
>>> def test_eigen_reconstruction(M3, eigenvalues, eigenvectors):
...     a = [[-0.49632839, -0.58968243, -0.68303647], [-0.89393102, -0.21497447, 0.46398207], [0.40824829, -0.81649658, 0.40824829]]
...     b = [[94.7493719, 0.0, 0.0], [0.0, -4.74937186, 0.0], [0.0, 0.0, -8.79571373e-16]]
...     M3_reconstructed = eigenvectors @ b @ a
...     assert np.allclose(M3_reconstructed, M3), f'Eigen reconstruction failed'
>>> test_eigen_reconstruction(M3, eigenvalues, eigenvectors)
```
